---
title: 'The basics of using R'
description: 'This chapter will serve as a gentle introduction to using R.'
---

## Simple calculations in R

```yaml
type: NormalExercise
key: 41d45aef32
lang: r
xp: 100
skills: 1
```

We will start out by running some simple calculations in R.

`@instructions`
- What is the result of calculating `abs(5)`?
- What is the result of calculating `abs(-5)`?
- Let's try to get help on the `abs()` function. Type `help(abs)` into the console.
- Try what happens if you type `?abs` instead.

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}
help(pi) #This displays the help for
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Good job! Now you know how to get help for a specific R command!")
```

---

## Getting help in R

```yaml
type: NormalExercise
key: 676ff8bf69
lang: r
xp: 100
skills: 1
```

The first thing you should know is how to get help in R.

`@instructions`
- What is the result of calculating `abs(5)`?
- What is the result of calculating `abs(-5)`?
- Let's try to get help on the `abs()` function. Type `help(abs)` into the console.
- Try what happens if you type `?abs` instead.

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}
help(pi) #This displays the help for
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Good job! Now you know how to get help for a specific R command!")
```
